<font size="4" style="color:red;"> **IMPORTANT: ** When submitting this homework notebook, please modify only the cells that start with:</font>

```python
# modify this cell
```

# 1. Die Rolls

Dice throws are another easy experiment to conduct. But first let us look at how we can reproduce the sampled data. To do this we can set the `seed` value of the random number generator to a chosen value.

```python
import matplotlib.pyplot as plt
import numpy as np
import ipywidgets as widgets
np.random.seed(666)
print(np.random.randint(9))
```

Output:
```
2
```

If we run the experiment again, we might get a different set of numbers.

```python
print(np.random.randint(9))
```

Output:
```
6
```

```python
np.random.seed(666)
print(np.random.randint(9))
print(np.random.randint(9))
```

Output:
```
2
6
```

## Exercise
Now consider a fair die. Each face has probability $\frac16$. We simulate $n$ die rolls and plot the empirical probability of each face, alongside the theoretical probability. 

```python
# modify this cell

# This is a decorator that creates the slider
@widgets.interact(n=(10,1000),continuos_update=False)
def probability_plot(n):
    """ 
        input: n (generate n random die rolls)
        output: Count (counting the occurance of each event )
    """
    #
    # YOUR CODE HERE
    #

    # plot the empirical values
    for i in range(6):
        Prob = Count[i,1:]/np.arange(1,n+1)
        plt.plot(np.arange(1, n + 1), Prob, linewidth=2.0, label='Face '+str(i+1))
    
    plt.plot(range(0, n), [1 / 6] * n, 'k', linewidth=3.0, label='Theoretical probability')
    plt.title("Empirical and theoretical probabilities of the 6 faces")
    plt.xlabel('Number of Iterations')
    plt.ylabel('Probability')
    plt.xlim([1, n])
    plt.ylim([0, 1])
    plt.legend()
    plt.show()
```

Next consider the event $E=\{2,4,6\}$ that the outcome is even. Clearly $P(E)=\frac36=\frac{1}{2}=0.5$.

The next cell simulates $n$ die rolls and plots the theoretical and empirical probabilities of $E$.

```python
# modify this cell

# This is a decorator that creates the slider
@widgets.interact(n=(10,10000),continuous_update=False)
def probability_plot_B(n):
    """ 
        input: n (generate n random die rolls)
        output: Prob_E (calculating the probability of even throw's)
        hint: counting the events of even numbers
    """
    #
    # YOUR CODE HERE
    #

    # plot
    plt.plot(range(1,n+1),Prob_E, 'b', linewidth= 2,label='Empirical probability')    
    plt.plot(range(1,n+1), [1 / 2] * n, 'k', linewidth= 2, label='Theoretical probability')
    
    plt.xlabel('Number of Iterations')
    plt.ylabel('Probability')
    plt.title("Odds of rolling an even number")
    plt.xlim([1, n])
    plt.ylim([0, 1])
    plt.legend()
    plt.show()
```

# 2. Conditional Probability and Baye's Rule


## Exercise

There are two urns $A$ and $B$. Urn $A$ contains $r_A$ red balls and $w_A$ white balls whereas urn $B$ contains $r_B$ red balls and $w_B$ white balls. One of the urns is picked at random and then one ball is picked at random from this urn. Write a function **conditional_probability** that calculates the conditional probability that the randomly chosen ball belonged to urn $A$ given that it is white. Assume that $\frac{r_A}{w_A}\neq\frac{r_B}{w_B}$.

<font  style="color:blue"> **Code**</font>
```python
rA, wA, rB, wB = 1., 2., 2., 1.
conditional__probability(rA, wA, rB, wB) 
```

<font  style="color:magenta"> **Output**</font>
```
0.6666666666666666
```

```python
# modify this cell

def conditional__probability(rA, wA, rB, wB):
    # inputs: all of them are of type 'float'
    # output: a variable of type 'float'
    
    #
    # YOUR CODE HERE
    #   

```

<!-- grade_id: ex1, points: 5, grade: true, solution: false -->
```python
assert( abs(conditional__probability(2., 4., 3., 3.) -0.5714285714285715) < 10**-5)  
assert( abs(conditional__probability(1., 3., 5., 2.) -0.7241379310344829) < 10**-5) 

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

# 3. Conditional probability analysis for math scores


Please try the same analysis for the students math scores following the same process conducted for Portuguese scores shown in Sec. 5 in "Experiment 2_Lecture" Notebook. **Insert new cells below to finish the process.**

```python
# modify this cell

data_math = pd.read_csv("student-mat.csv")
```